# Lab: GPA

In this lab, you will explore Wade's famous GPA dataset and find out about the easiest, hardest, and most popular courses at Illinois!

The dataset you are using is this lab is actually the same dataset that Prof Wade used to make the GPA visualizations :) Check this out here if you haven't seen it already: [https://waf.cs.illinois.edu/discovery/grade_disparity_between_sections_at_uiuc/](https://waf.cs.illinois.edu/discovery/grade_disparity_between_sections_at_uiuc/)

## Group

Edit the next Python cell to add information about who you're working within your lab section:

```python
# First, meet your CAs and TA if you haven't already!
# ...first name is enough, we'll know who they are! :)
ta_name = "Daniel"
ca1_name = "Tanvi"
ca2_name = "Kaitlyn"


# Also, make sure to meet your team for this lab! Find out their name, what major they're in,
# and learn something new about them that you never knew before!
partner1_name = "Jack"
partner1_netid = "johnsm4"
partner1_major = "Statistics"

partner2_name = "Degaulle"
partner2_netid = "ddai9"
partner2_major = "Food Micro Biology"

partner3_name = "Jay"
partner3_netid = "jayaw2"
partner3_major = "CS + Math"
```

## Table of Contents <a id='toc'></a>

**Use the Outline view (View > Command Palette > Type "outline") for navigation.**

* Puzzle 1
    * Puzzle 1.1
    * Puzzle 1.2
* Puzzle 2
    * Puzzle 2.1
    * Puzzle 2.2
    * Puzzle 2.3 
    * Puzzle 2.4
* Puzzle 3
    * Puzzle 3.1
    * Puzzle 3.2
    * Puzzle 3.3
    * Puzzle 3.4
    * Puzzle 3.5

## Puzzle 1: Basic Data Manipulation

### Puzzle 1.1: Loading the Dataset

For the first part of our lab, we'll be using the GPA dataset, which should be automatically downloaded when you fetch this notebook. 

In the following cell, load the pandas library and read the dataset into `df`.

<!-- grade_id: cell-51fef9428e349da7, points: 0, grade: false, solution: true -->
```python
#import the gpa.csv dataset below
import pandas as pd
df = pd.read_csv("/Users/saber/Desktop/stat107/srahm3/lab_gpa/gpa.csv")
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb Cell 4' in <module>
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000003?line=0'>1</a> #import the gpa.csv dataset below
----> <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000003?line=1'>2</a> import pandas as pd
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000003?line=2'>3</a> df = pd.read_csv("/Users/saber/Desktop/stat107/srahm3/lab_gpa/gpa.csv")

ModuleNotFoundError: No module named 'pandas'
```

### Puzzle 1.2: Adding Helper Columns

Add two new columns, `"Total Students"` and `"Average GPA"`, that compute the total number of students in each course and the average GPA for each course, respectively.

You should not include students with a `"W"` in either column. The grade point scale can be found at https://registrar.illinois.edu/courses-grades/explanation-of-grades/.

Hint (double-click this cell to see it): <span style="color:#ffffff00">it is not necessary but you can try the dot() function; see https://stackoverflow.com/questions/18419962/how-to-compute-weighted-sum-of-all-elements-in-a-row-in-pandas for some examples.</span>

<!-- grade_id: cell-1043285927ac8a85, points: 0, grade: false, solution: true -->
```python
weight = [4, 4, 3.67, 3.33, 3, 2.67, 2.33, 2, 1.67, 1.33, 1, 0.67, 0]
grade = ["A+", "A", "A-", "B+", "B", "B-", "C+", "C", "C-", "D+", "D", "D-", "F"]
df["Total Students"] = df[grade].sum(axis=1)
df["Average GPA"] =  df[grade].dot(weight) / df["Total Students"]
df
```

Output:
```
NameError                                 Traceback (most recent call last)
/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb Cell 6' in <module>
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000005?line=0'>1</a> weight = [4, 4, 3.67, 3.33, 3, 2.67, 2.33, 2, 1.67, 1.33, 1, 0.67, 0]
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000005?line=1'>2</a> grade = ["A+", "A", "A-", "B+", "B", "B-", "C+", "C", "C-", "D+", "D", "D-", "F"]
----> <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000005?line=2'>3</a> df["Total Students"] = df[grade].sum(axis=1)
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000005?line=3'>4</a> df["Average GPA"] =  df[grade].dot(weight) / df["Total Students"]
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000005?line=4'
```

<!-- grade_id: cell-46729cc74c213e52, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1.2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, and RE-RUN your code and then this cell.
assert(df[
  (df["YearTerm"] == "2010-su") & \
  (df["Subject"] == "STAT") & \
  (df["Number"] == 410) \
]["Total Students"].sum() == 31), "It appears that your \"Total Students\" column is incorrect."
assert(df[
  (df["YearTerm"] == "2018-fa") & \
  (df["Subject"] == "CS") & \
  (df["Number"] == 225) \
]["Total Students"].sum() == 697), "It appears that your \"Total Students\" column is incorrect."
assert( abs(df[
  (df["YearTerm"] == "2018-fa") & \
  (df["Subject"] == "STAT") & \
  (df["Number"] == 100) \
]["Average GPA"].mean() - 3.247) < 0.01), "It appears that your \"Average GPA\" column is incorrect."
## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

Output:
```
NameError                                 Traceback (most recent call last)
/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb Cell 7' in <module>
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000006?line=0'>1</a> ## == TEST CASES for Puzzle 1.2 ==
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000006?line=1'>2</a> # - This read-only cell contains test cases for your previous cell.
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000006?line=2'>3</a> # - If this cell runs without any error our output, you PASSED all test cases!
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000006?line=3'>4</a> # - If this cell results in any errors, check your previous cell, make changes, and RE-RUN your code and then this cell.
----> <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3
```

## Puzzle 2: Finding the Easiest and Hardest Courses

### Puzzle 2.1: Hardest Courses in the Dataset

Based on the `"Average GPA"` column, find the **30 hardest courses** at Illinois, i.e., the 30 courses with the lowest "Average GPA". Store the results in `df_hard`.

Hint (double-click this cell to see it): <span style="color:#ffffff00">use the nsmallest() function</span>

<!-- grade_id: cell-13f6d92b7a658526, points: 0, grade: false, solution: true -->
```python
df_hard = df.nsmallest(30, "Average GPA")
df_hard
```

Output:
```
NameError                                 Traceback (most recent call last)
/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb Cell 9' in <module>
----> <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000008?line=0'>1</a> df_hard = df.nsmallest(30, "Average GPA")
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000008?line=1'>2</a> df_hard

NameError: name 'df' is not defined
```

### Puzzle 2.2: "Level" of Hardest Courses

Find the average `"Number"` of the **30 hardest courses** at Illinois. Store the result as `df_hard_avg`.

<!-- grade_id: cell-a7106dca7694711c, points: 0, grade: false, solution: true -->
```python
df_hard_avg = df_hard["Number"].mean()
df_hard_avg
```

Output:
```
NameError                                 Traceback (most recent call last)
/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb Cell 11' in <module>
----> <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000010?line=0'>1</a> df_hard_avg = df_hard["Number"].mean()
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000010?line=1'>2</a> df_hard_avg

NameError: name 'df_hard' is not defined
```

<!-- grade_id: cell-46729cc74c213e52, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2.1 and 2.2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check your previous cell(s), make changes, and RE-RUN your code and then this cell.
import math
assert( math.isclose(df_hard_avg, 4988/30) ), "It appears that your average \"Number\" is incorrect. Please check both Puzzle 2.1 and Puzzle 2.2."
## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

Output:
```
NameError                                 Traceback (most recent call last)
/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb Cell 12' in <module>
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000011?line=0'>1</a> ## == TEST CASES for Puzzle 2.1 and 2.2 ==
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000011?line=1'>2</a> # - This read-only cell contains test cases for your previous cell.
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000011?line=2'>3</a> # - If this cell runs without any error our output, you PASSED all test cases!
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000011?line=3'>4</a> # - If this cell results in any errors, check your previous cell(s), make changes, and RE-RUN your code and then this cell.
      <a href='vscode-notebook-cell:/Users/saber/Desktop/st
```

### Puzzle 2.3: ❓ Individual Reflection Question ❓ Are the Hardest Courses Really "Hard"?

After solving Puzzle 2.1 and 2.2, Joe has the following claim:

"We know that the undergraduate courses are coded from 1 to 499, where a larger number (in the hundreds place) usually implies more advanced materials. Based on Puzzle 2.2, the results in Puzzle 2.1 clearly do not represent the hardest courses at Illinois."

Comment on Joe's claim below.

I think that 2.1 shows that people that try classes in subjects they are interested in but have no experience do have lower gpas, so it lower than it should be, like Joe said. That also likely means that people in the upper level classes have a good foundation of what they should know so they can pass that class.

### Puzzle 2.4: Easiest Courses in 2020

Based on the `"Average GPA"` column, find the **30 easiest courses in 2020** at Illinois. Store the results in `df_easy_2020`.

Hint (double-click this cell to see it): <span style="color:#ffffff00">use the nlargest() function</span>

<!-- grade_id: cell-44d17a2ef57b1a19, points: 0, grade: false, solution: true -->
```python
df_easy_2020 = df[df["Year"] == 2020].nlargest(30, "Average GPA")
df_easy_2020
```

Output:
```
NameError                                 Traceback (most recent call last)
/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb Cell 16' in <module>
----> <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000015?line=0'>1</a> df_easy_2020 = df[df["Year"] == 2020].nlargest(30, "Average GPA")
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000015?line=1'>2</a> df_easy_2020

NameError: name 'df' is not defined
```

<!-- grade_id: cell-af77ae707b9fa776, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2.4 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error or output, you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, RE-RUN your code, and then RE-RUN this cell.
import math
assert( len(df_easy_2020) == 30 ), "You should only have the 30 easiest courses in 2020."
assert( math.isclose(df_easy_2020["Average GPA"].mean(), 3.9730388348773817) ), "Are you sure you have the 30 easiest courses in 2020?"
## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

Output:
```
NameError                                 Traceback (most recent call last)
/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb Cell 17' in <module>
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000016?line=0'>1</a> ## == TEST CASES for Puzzle 2.4 ==
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000016?line=1'>2</a> # - This read-only cell contains test cases for your previous cell.
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000016?line=2'>3</a> # - If this cell runs without any error or output, you PASSED all test cases!
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000016?line=3'>4</a> # - If this cell results in any errors, check your previous cell, make changes, RE-RUN your code, and then RE-RUN this cell.
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/sr
```

## Puzzle 3: Finding Data by Subject

### Puzzle 3.1: Total Number of Grades by Subject

Find the total number of each grade (eg: "A+"s, "A"s, "A-"s, etc.) given by different subjects. Store the results in `df_group`.

Hint (double-click this cell to see it): <span style="color:#ffffff00">use the groupby() and agg() function</span>

<!-- grade_id: cell-7d31cf501bfd8025, points: 0, grade: false, solution: true -->
```python
df_group = df.groupby("Subject").agg("sum").reset_index()
df_group
```

Output:
```
NameError                                 Traceback (most recent call last)
/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb Cell 19' in <module>
----> <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000018?line=0'>1</a> df_group = df.groupby("Subject").agg("sum").reset_index()
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000018?line=1'>2</a> df_group

NameError: name 'df' is not defined
```

<!-- grade_id: cell-e9ad64dffa5a1cf1, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 3.1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, RE-RUN your code, and then RE-RUN this cell.
assert( len(df_group) == 163 ), "Are you sure you grouped by subjects?"
assert( df_group[ df_group["Subject"] == "CS" ]["A"].values[0] == 45290 ), "Are you sure you counted the right number of each subject?"
## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

Output:
```
NameError                                 Traceback (most recent call last)
/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb Cell 20' in <module>
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000019?line=0'>1</a> ## == TEST CASES for Puzzle 3.1 ==
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000019?line=1'>2</a> # - This read-only cell contains test cases for your previous cell.
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000019?line=2'>3</a> # - If this cell runs without any error our output, you PASSED all test cases!
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000019?line=3'>4</a> # - If this cell results in any errors, check your previous cell, make changes, RE-RUN your code, and then RE-RUN this cell.
----> <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/s
```

### Puzzle 3.2: Average GPA by Subject

In `df_group`, add a new column called `"Overall Average GPA"` that computes the overall GPA of each subject at Illinois.

<!-- grade_id: cell-7d3923eba40e1f45, points: 0, grade: false, solution: true -->
```python
df_group["Overall Average GPA"] = df_group[grade].dot(weight) / df_group["Total Students"]
df_group
```

Output:
```
NameError                                 Traceback (most recent call last)
/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb Cell 22' in <module>
----> <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000021?line=0'>1</a> df_group["Overall Average GPA"] = df_group[grade].dot(weight) / df_group["Total Students"]
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000021?line=1'>2</a> df_group

NameError: name 'df_group' is not defined
```

<!-- grade_id: cell-81b547cf0d42b67e, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 3.2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error or output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, RE-RUN your code, and then RE-RUN this cell.
import math
assert( math.isclose(df_group[df_group["Subject"] == "CS"]["Overall Average GPA"].values[0], 3.2788726192300266) ), "Your overall average GPA does not appear to be correct."
## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

Output:
```
NameError                                 Traceback (most recent call last)
/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb Cell 23' in <module>
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000022?line=0'>1</a> ## == TEST CASES for Puzzle 3.2 ==
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000022?line=1'>2</a> # - This read-only cell contains test cases for your previous cell.
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000022?line=2'>3</a> # - If this cell runs without any error or output, you PASSED all test cases!
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000022?line=3'>4</a> # - If this cell results in any errors, check you previous cell, make changes, RE-RUN your code, and then RE-RUN this cell.
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/sra
```

### Puzzle 3.3: Visualizing the Grades of Subject Equal to `"STAT"`

Plot a density histogram of the grades (`"A+"` to `"F"`) for the subject `"STAT"`.

Hint (double-click this cell to see it): 
<span style="color:#ffffff00">
    First, extract the row where the subject is "STAT".
    Then, extract the columns "A+" to "F" and divide them by the total number of students.
    Finally, use plot.bar().
</span>

```python
dfstat = df_group[df_group["Subject"] == "STAT"]
dfc = dfstat[grade] / dfstat["Total Students"].values[0]
dfc.plot.bar()
```

Output:
```
NameError                                 Traceback (most recent call last)
/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb Cell 25' in <module>
----> <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000024?line=0'>1</a> dfstat = df_group[df_group["Subject"] == "STAT"]
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000024?line=1'>2</a> dfstat[grade]
      <a href='vscode-notebook-cell:/Users/saber/Desktop/stat107/srahm3/lab_gpa/lab_gpa.ipynb#ch0000024?line=2'>3</a> dfc = dfstat[grade] / dfstat["Total Students"].values[0]

NameError: name 'df_group' is not defined
```

### Puzzle 3.4: ❓ Individual Reflection Question ❓ 

Density Histogram vs. Frequency Histogram

Why should we use density histogram in Puzzle 3.3? Is it meaningful to compare the frequency histogram between different subjects?

The density histogram makes it easy to identify what grade was most common between all the students where as frequency would just show where grade differences spiked.

### Puzzle 3.5: Find Your Major

Finally, find your major/subject in the dataset and check out how hard it is! 

Feel free to look around and explore! Part of being a data scientist is being creative and using Python to answer questions that you have.  Write your code below and write comments in Puzzle 3.6 to explain what your code does. For example, what are the hardest classes in your major that have had at least 50 students?

<!-- grade_id: cell-116a2bb23ee81616, points: 0, grade: false, solution: true -->
```python
# Feel free to explore. Credit will be given as long as you try some new code (i.e., not directly copied from any puzzles).

dis = df[df["Subject"] == "IS"]
dis = dis.nsmallest(10, "Average GPA")
dis
```

Output:
```
Year    Term YearTerm Subject  Number                    Course Title  \
26573  2015  Spring  2015-sp      IS     502      Libraries Info and Society   
13968  2017    Fall  2017-fa      IS     490       Introduction to Databases   
52900  2010  Spring  2010-sp      IS     310     Computing in the Humanities   
32560  2014  Summer  2014-su      IS     502      Libraries Info and Society   
27530  2015  Summer  2015-su      IS     502      Libraries Info and Society   
32561  2014  Summer  2014-su      IS     502      Libraries Info and Society   
27529  2015  Summer  2015-su      IS     502      Libraries Info and Society   
8903   2018    Fall  2018-fa      IS     542                Data, Stat, Info   
17398  2017  Summer  2017-su      IS     507  Intr to Bibliographic Metadata   
8904   2018    Fall  2018-fa      IS     542                Data, Stat, Info   

       A+  A  A-  B+  ...  C  C-  D+  D  D-  F  W      Primary Instructor  \
26573   0  2   1  16  ...  1   0   0  1   0  0  0
```

### Puzzle 3.6: ❓ Individual Reflection Question ❓ 

Explain what you did in Puzzle 3.5 here, as well as what you found!  This is your chance to interpret your results and teach us something that you found interesting :)

I pulled all the courses that were under my major of Information Science and pulled out the lowest gpas from all the classes. I found it interesting that the lowest GPA was still above a 3.

# Finale!

You're almost done -- congratulations!

You need to do two more things:

1.  Save your work. To do this, go to File -> Save All

2.  After you have saved, exit this notebook and follow the webpage instructions to commit this lab to your Git repository!